---
title       : Programming basics
description : Main programming constructs in Python

--- type:NormalExercise lang:python xp:100 skills:2 key:2a7a755065
## Demonstration of error
If the explanation contains inequalities, the output may be wrong: $x<3$ and $y>2$

*** =instructions
In the instruction block this does not work: $x<3$ and $y>2$

*** =hint
In the hint block this does not work: $x<3$ and $y>2$

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
